# Setup
CS1302 Introduction to Computer Programming
___

## JupyterHub

### How to access the JupyterHub server?

1. Enter the url of the Jupyterhub server [divedeep.cs.cityu.edu.hk](https://divedeep.cs.cityu.edu.hk) in a web browser.
2. Enter your [EID](https://www.cityu.edu.hk/esu/eid.htm) and Password in the fields `Username` and `Password` respectively.
3. Click the `Sign in` button.

![log in](images/login.dio.svg)

4. Choose `CS1302` as the server option and click `Start`.  
  N.b., `Minimal` is only for trouble-shooting and is insufficient for running the course materials.

![server options](images/serveroptions.dio.svg)

5. The JupyterLab interface should appear.

![jupyterlab](images/jupyterlab.dio.svg) 

  - In case or error, you can restart your server using the `Hub Control Panel` under the `File` menu (shown above) or visit:  
    https://divedeep.cs.cityu.edu.hk/hub/home  
    You can also refresh the page, re-login, or clear your browser cache to reset your browser.
  - To submit homework assignment, you will need to `Launch Classic Notebook` interface under the `Help` menu (shown above) or visit:  
    https://divedeep.cs.cityu.edu.hk/user-redirect/tree  
    or replace `lab` in your browser url (illustrated above in yellow) by `tree` can press `Enter`.  

## Jupyter Notebook

### How to access course materials?

You can click on the individual links to the lecture/lab notebooks on the [course homepage](https://canvas.cityu.edu.hk/courses/42868) or visit  
https://divedeep.cs.cityu.edu.hk/user-redirect/git-pull?repo=https%3A%2F%2Fgithub.com%2Fccha23%2Fcs1302&urlpath=lab%2Ftree%2F%2Fcs1302%2F&branch=main 


The links all go to the JupyterHub server and [*git-pull*](https://git-scm.com/docs/git-pull) the released notebooks to the `cs1302` folder under your home directory.

### How to complete a lab assignment?

After opening a notebook:

1. Click `Help->Welcome Tour` to learn the JupyterLab interface. 
1. Click `Help->Notebook Tour` to learn about the jupyter notebook.

**Exercise** In learning a new programming language, the first program to write is often the ["Hello, World!"](https://en.wikipedia.org/wiki/%22Hello,_World!%22_program) program, which says Hello to the world. Type the program `print('Hello, World!')` below and run it with `Shift+Enter`.

<!-- grade_id: helloworld, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

We often ask you to write a code in a particular cell. Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE".

In order to check your work thoroughly, there will visible and hidden test cases. The following is a visible test you can run to check your answer: The test returns an assertion error only if your program does not print the correct message.

<!-- grade_id: test-helloworld, points: 1, grade: true, solution: false -->
```python
# Run this test cell right after running your "Hello, World!" program.
import sys, io
old_stdout, sys.stdout = sys.stdout, io.StringIO()
exec(In[-2])
printed = sys.stdout.getvalue()
sys.stdout = old_stdout
assert printed == 'Hello, World!\n'
```

**Tips**
1. You can repeatedly modify your solution and run the test cell until your solution passes the test. You are not required to know how the test cell is written. 
1. To assess your solution thoroughly, we often run new tests hidden from you after you have submitted your notebook. There is no partial credit for a partially correct solution that works for the visible test but fails for the hidden test. Therefore, *you should ensure your solution works in general rather than just the visible tests*.
1. If you open the same notebook multiple times in different browser windows, be careful in making changes in different windows. Inconsistent changes may lead to conflicts or loss of your data.
1. If your notebook fails to run any code, the Kernel might have died. You can restart the kernel with `Kernel->Restart`. If restarting fails, check your code cells to see if there is any code that breaks the kernel.

### How to submit a notebook

- Although the first Lab does not count towards your final grade, you are required to submit it, to get familiar with the procedure.
- Before you submit, make sure everything runs as expected:
    1. **Restart the kernel**: `Kernel->Restart Kernel...` 
    1. **run all cells**: `Run->Run All Cells`

To submit your notebook:
1. Start the classic notebook interface by choosing `Help->Launch Classic Notebook`.
1. Go to `Assignment` tab of JupyterHub where you fetched the Lab assignment. 
1. Expand the Lab folder and click the `validate` button next to the notebook(s) to check if all the visible tests pass.
1. Click `Submit` to submit your notebook. 
1. You may submit as many times as you wish before the due date as we collect your latest submission for grading. 
    - *No late submission* will be collected without valid justifications.
    - *Double check* that you have submitted the correct Lab assignment.
    - You are responsible for *recording your submission attempt* with a valid timestamp in case of technical issues.

**Tips**
1. You normally have at least 5 days to work on the lab after your lab session. 
1. You can check the due dates of all the labs from the course homepage. 
1. You may seek help from us or your classmates. However, you must write your own solution and indicate who your collaborators are using the code:

```python
COLLABORATORS = ['WONG Xiu Fong', 'LEE Man Kit']
```