# Improved Quadratic Equation Solver

**CS1302 Introduction to Computer Programming**
___

In this notebook, we will improve the quadratic equation solver in the previous lab using conditional executions.  
First of all, run the following to setup the environment.

```python
from ipywidgets import interact
import math
import numpy as np
```

## Discriminant

Recall that a quadratic equation is

$$
ax^2+bx+c=0
$$

where $a$, $b$, and $c$ are real-valued coefficients, and $x$ is the unknown variable.

---

**Definition**

The disciminant of a quadratic equation is defined as

$$
\Delta := b^2 - 4ac.
$$ (discriminant)

---

The discriminant $\Delta$ discriminates the roots $\frac{-b\pm \sqrt{\Delta}}{2a}$ of a quadratic equation:

---

**Proposition**

The roots of a quadratic equation {eq}`quadratic` are the same (repeated) equal to $-\frac{b}{2a}$ when the discriminant is zero, i.e., $\Delta=0$.

---

**Exercise** Complete the following function by assigning `roots` only one root when the discriminant is zero. E.g., if $(a,b,c)=(1,-2,1)$, then `roots` should be assigned the value `1.0` instead of `1.0, 1.0`.

<!-- grade_id: zero_determininant, points: 0, grade: false, solution: true -->
```python
def get_roots(a, b, c):
    d = b ** 2 - 4 * a * c  # discriminant
    if math.isclose(d, 0):
        # YOUR CODE HERE
        raise NotImplementedError()
    return roots
```

---

**Hint**

Consider using the [if statement](https://docs.python.org/3/reference/compound_stmts.html#if) as follows:

```Python
def get_roots(a, b, c):
    d = b ** 2 - 4 * a * c  # discriminant
    if math.isclose(d, 0):
        roots = ___  # repeated root
    else:
        d **= 0.5
        roots = ___, ___
    return roots
```

This is a solution template where `___` needs to be filled in with appropriate code. If you are a more efficient/simpler implementation, you need *not* follow the solution template. E.g., to challenge your understanding of short-circuit evaluation, you can write your solution using boolean operations without any if statement.

---

<!-- grade_id: test-zero_determinant, points: 1, grade: true, solution: false -->
```python
# tests
def test_get_roots(roots, a, b, c):
    def mysort(c):
        return c.real, c.imag

    roots_ = get_roots(a, b, c)
    assert (
        hasattr(roots, "__len__")
        and np.isclose(
            sorted(list(roots), key=mysort), sorted(list(roots_), key=mysort)
        ).all()
        or roots == roots_ or np.isclose(roots, roots_)
    )


test_get_roots((-1.0, 0.0), 1, 1, 0)
test_get_roots(0.0, 1, 0, 0)
```

<!-- grade_id: hidden_test-zero_determinant, points: 1, grade: true, solution: false -->
```python
# hidden tests
```

**Exercise** Why use `math.isclose(d, 0)` instead of `d == 0`?

<!-- grade_id: isclose, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

## Linear equation

If $a=0$, the earlier formula for the roots are invalid due to division by zero. Nevertheless, the equation remains valid:

$$
bx + c=0.
$$

**Exercise** Improve the function `get_roots` to return the root $-\frac{c}{b}$ when $a=0$.

<!-- grade_id: linear, points: 0, grade: false, solution: true -->
```python
def get_roots(a, b, c):
    d = b ** 2 - 4 * a * c
    # YOUR CODE HERE
    raise NotImplementedError()
    return roots
```

---

**Hint** 

Solution template:
```Python
def get_roots(a, b, c):
    d = b ** 2 - 4 * a * c    # discriminant
    if ___:
        roots = ___
    elif math.isclose(d, 0):
        roots = ___  # repeated root
    else:
        d **= 0.5
        roots = ___
    return roots
```

---

<!-- grade_id: test-linear, points: 1, grade: true, solution: false -->
```python
# tests
test_get_roots((-1.0, -0.0), 1, 1, 0)
test_get_roots(0.0, 1, 0, 0)
test_get_roots(0.5, 0, -2, 1)
```

<!-- grade_id: hidden_test-linear, points: 1, grade: true, solution: false -->
```python
# hidden tests
```

## Degenerate cases

What if $a=b=0$? In that case, the equation becomes

$$
c = 0
$$
which is always satisfied if $c=0$ but never satisfied if $c\neq 0$. 

**Exercise** Improve the function `get_roots` to return root(s) under all cases:
- If $a=0$ and $b\neq 0$, assign `roots` to the single root $-\frac{c}{b}$. 
- If $a=b=0$ and $c\neq 0$, assign `roots` to `None`.  
    Note that `None` is an object, not a string.
- If $a=b=c=0$, there are infinitely many roots. Assign to `roots` the tuple `-float('inf'), float('inf')`.  
    Note that `float('inf')` converts the string `'inf'` to a floating point value that represents $\infty$.

<!-- grade_id: degenerate, points: 0, grade: false, solution: true -->
```python
def get_roots(a, b, c):
    d = b ** 2 - 4 * a * c
    # YOUR CODE HERE
    raise NotImplementedError()
    return roots
```

---

**Hint** 

Use nested `if` statements as in the solution template:

```Python
def get_roots(a, b, c):
    d = b**2 - 4 * a * c
    if ___:
        if ___:
            if ___:
                roots = -float('inf'), float('inf')
            else:
                roots = None
        else:
            ___
    elif math.isclose(d, 0):
        roots = ___  # repeated root
    else:
        d **= 0.5
        roots = ___
    return roots
```

---

<!-- grade_id: test-degenerate, points: 1, grade: true, solution: false -->
```python
# tests
test_get_roots((-1.0, 0.0), 1, 1, 0)
test_get_roots(0.0, 1, 0, 0)
test_get_roots((-float('inf'), float('inf')), 0, 0, 0)
test_get_roots(None, 0, 0, 1)
test_get_roots(0.5, 0, -2, 1)
test_get_roots(1.0, 1, -2, 1)
```

<!-- grade_id: test_test-degenerate, points: 1, grade: true, solution: false -->
```python
# hidden tests
```

After you have complete the exercises, you can run your robust solver below:

```python
# quadratic equations solver
@interact(a=(-10,10,1),b=(-10,10,1),c=(-10,10,1))
def quadratic_equation_solver(a=1,b=2,c=1):
    print('Root(s):',get_roots(a,b,c))
```

## Flowcharts (optional)

Let's play a game to test your understanding of boolean expressions. Run the following program and choose your input so that 1, 2, 3 are each printed in a separate line as follows:
``` 
1
2
3
```

```python
input(1) or input(2) and input(3)
```

---

**Hint**

The following flowchart describes the [short-circuit evaluation](https://docs.python.org/3/reference/expressions.html?highlight=precedence#boolean-operations) for the boolean expression as follows:

![1or2and3](images/1or2and3.svg)

---

**How to read the flowchart?**

A flowchart uses arrows to connects a set of [annotated blocks][bb]. The rules were first specified by ANSI and later adopted in [ISO 5807][iso].

[bb]: https://en.wikipedia.org/wiki/Flowchart#Building_blocks
[iso]: https://webstore.ansi.org/Standards/ISO/ISO58071985

For the above game, start from the first decision block at the top and choose the input appropriately to transition along the `yes` arrows.

**Why use a program flowchart?**

A program flowchart is a powerful way of describing an algorithm quickly: Unlike a text-based programming language, 
- The rules governing the program flow can be shown explicitly by arrows.
- The annotated graphical blocks can convey the meaning faster using both visual clues. 

A programmer can draft a program in flowchart without worrying about syntax (indentation, colon, ...). E.g., the above flow chart can be implement as follows using the if statement instead of the boolean operators.

```python
if not input(1):
    if input(2):
        input(3)
```

**How to convert the program to a flowchart automatically?**

We can use [PyFlowchart](https://github.com/cdfmlr/pyflowchart):

```python
from pyflowchart import Flowchart

def code2flow(code):
    '''
    Given input code (as a string), returns the flowchart (as a string) that can be rendered by
    http://flowchart.js.org
    '''
    return Flowchart.from_code(code, simplify=False).flowchart()
```

Pass the code as a string to `code2flow`:

```python
fc = code2flow('''
if not input(1):
    if input(2):
        input(3)
''')
print(fc)
```

The output is a string in a [domain-specific language (DSL)][dsl] that can be [rendered here with flowchart.js][fc]. Simply overwrite a demo textbox there with the printed flowchart code above.

[dsl]: https://en.wikipedia.org/wiki/Domain-specific_language
[fc]: http://flowchart.js.org

**How to write a documentation with flowcharts?  
What about other DSL for special contents such as diagrams, equations, lists, ...?**

[Markdown](https://en.wikipedia.org/wiki/Markdown) is becoming popular in writing documentations. jupyter notebooks, jupyterbook, github, ... all supports markdown, but to different extents:

- Jupyter notebook such as this one supports [CommonMark](https://commonmark.org).
- Our [cs1302 jupyterbook](https://www.cs.cityu.edu.hk/~ccha23/cs1302book) is written in [MyST](https://myst-parser.readthedocs.io/en/latest/syntax/syntax.html) markdown.
- Github uses [Github Flavored Markdown (GFM).](https://github.github.com/gfm/)

Unfortunately, none of them support the flowchart code currently.

Nevertheless, many editors such as [Typora](https://typora.io) supports advanced Markdown language with live preview. We can also extend VS Code to do the same using [Markdown Preview Enhanced](https://shd101wyy.github.io/markdown-preview-enhanced/#/) as follows:

---

**Tip**

You can install any extension on the VS Code interface of the JupyterHub server. To render flowchart code:

1. Open `Launcher` (click on the **+** sign in the top left or press `Shift + Ctrl + L`) and click on `VS Code` to create a VSCode tab
2. Click the extension tab on the left and search for *Markdown Preview Enhanced*.
3. Click install and refresh the browser.
4. Click the file explorer tab and create a new file say `test.md` file.
5. Insert the flowchart in DSL to the `test.md` [as follows](https://shd101wyy.github.io/markdown-preview-enhanced/#/diagrams?id=flow-charts):  
    ````Markdown
    ```flow
    cond3=>condition: if (not input(1))
    cond8=>condition: if input(2)
    sub12=>subroutine: input(3)

    cond3(yes)->cond8
    cond8(yes)->sub12
    ```
    ````
6. To show the live preview, click the `open preview` button of Markdown Preview Enhanced located on the tool bar at the top right-hand corner.

---